rule_files:
  - ../rules-glbc.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: up{}
        values: "1+0x5 0+0x15"
    alert_rule_test:
      - eval_time: 5m
        alertname: ${ALERT_NAME}
        exp_alerts: []
      - eval_time: 20m
        alertname: ${ALERT_NAME}
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: 'Summary of the alert'
              description: 'More detailed description of the alert and cause of firing'
              runbook_url: '${RUNBOOK_PATH}${ALERT_NAME}.adoc'